# Homework #10

## PHYS 2600

__Important notice:__ All cells in your notebook will be run, start to finish, using a fresh kernel when it is graded! To make sure the graded notebook looks like what you expect, we recommend selecting "Runtime > Restart session and run all" from the menu above in Colab before you finish.

```python
%matplotlib inline

import matplotlib.pyplot as plt
import numpy as np
```

If you worked collaboratively on this assignment, __include the names of your collaborators in the cell below:__

###Just Me

## 10.0 - Homework correction (3 points)

In the cell below, explain and correct __one mistake__ you made on homework #8.  If you got a perfect score, great, tell us what your favorite problem was, or use the space to give any other feedback you might have on the class/tutorials/homework.

<!-- grade_id: hw10-correction, points: 3, grade: true, solution: true -->
### Great class The only thing tha I'm having trouble at is the Physics & math not the programming

## 10.1 - Fun and games in Python (20 points)


### Part A (3 points)

Below is a list containing each card in a standard deck of 52 playing cards.  Each card has a __rank__, ranging from the numbers 2-10 or the "royal" values J/Q/K/A, and a __suit__, one of spades/clubs/hearts/diamonds (♠︎/♣︎/♥︎/♦︎).

__Write a function `deal_hand(n)`__ that will produce a list of $n$ cards, _drawn without replacement_ from the `deck` list.  (_Hint: your function can easily be one line long!_)

<!-- grade_id: cell-f46a05984a8267d6, points: 0, grade: false, solution: true -->
```python
deck = [
    "2♠",
    "3♠",
    "4♠",
    "5♠",
    "6♠",
    "7♠",
    "8♠",
    "9♠",
    "10♠",
    "J♠",
    "Q♠",
    "K♠",
    "A♠",
    "2♣",
    "3♣",
    "4♣",
    "5♣",
    "6♣",
    "7♣",
    "8♣",
    "9♣",
    "10♣",
    "J♣",
    "Q♣",
    "K♣",
    "A♣",
    "2♥",
    "3♥",
    "4♥",
    "5♥",
    "6♥",
    "7♥",
    "8♥",
    "9♥",
    "10♥",
    "J♥",
    "Q♥",
    "K♥",
    "A♥",
    "2♦",
    "3♦",
    "4♦",
    "5♦",
    "6♦",
    "7♦",
    "8♦",
    "9♦",
    "10♦",
    "J♦",
    "Q♦",
    "K♦",
    "A♦",
]


def deal_hand(n):
    #
    return np.random.choice(deck, n, replace=False) # I liked the on line idea

```

<!-- grade_id: hw10_1a, points: 3, grade: true, solution: false -->
```python
test_hand = deal_hand(5)
print(test_hand)
print(len(test_hand[0]))

assert len(test_hand) == 5
for card in test_hand:
    assert card in deck

whole_deck = deal_hand(52)
assert set(whole_deck) == set(deck)
```

Output:
```
['7♣' 'Q♥' 'Q♦' 'J♥' '6♦']
2
```

### Part B (10 points)

Now we'd like to recognize whether a hand contains a certain number of matching ranks or suits.  Implement the following functions:

- `has_4kind(hand)` should return `True` if a hand has 4 cards with the same rank ("three-of-a-kind"), and `False` otherwise.
- `has_flush(hand)` should return `True` if a hand has 5 cards with the same suit (a "flush"), and `False` otherwise.

__(Make sure your code deals properly with `10`, which has an extra character compared to the other ranks!)__

_(Hint: you should start by splitting each card into a rank and a suit.  Remember, you can treat strings like lists and uses slices and indexing on them.)_


_(Another hint: as we've seen, `np.unique` is very useful for counting how many instances there are of all discrete values in a list.)_



<!-- grade_id: cell-1588122b761d3eb0, points: 0, grade: false, solution: true -->
```python
# I need to extract the rank of the hand first
# I will create a helper function - I know there's a better way but my brainn
# is froze at the moment

def extract_rank(card):
  return card[:-1] # a better way
  # if card.startswith("10"):
  #   return "10"
  # else:
  #   return card[0];

def has_4kind(hand):
  ranks = [extract_rank(card) for card in hand];
  value,count = np.unique(ranks,return_counts=True)
  return np.any(count >= 4)
#
```

<!-- grade_id: hw10_1b_1, points: 5, grade: true, solution: false -->
```python
test_hand_1 = ["8♠", "3♣", "8♥", "2♦", "8♦"]
test_hand_2 = ["K♠", "K♣", "K♥", "K♦", "8♦"]
test_hand_3 = ["10♠", "10♣", "10♥", "K♦", "2♦"]
test_hand_4 = ["3♥", "A♥", "K♥", "7♥", "8♥"]

print(has_4kind(test_hand_1))
print(has_4kind(test_hand_2))

assert not has_4kind(test_hand_1)
assert has_4kind(test_hand_2)
assert not has_4kind(test_hand_3)
assert not has_4kind(test_hand_4)

test_hand_5 = ["10♣", "3♣", "K♣", "Q♣", "J♣"]
test_hand_6 = ["8♥", "K♠", "K♣", "K♥", "K♦"]

assert not has_4kind(test_hand_5)
assert has_4kind(test_hand_6)
assert has_4kind(deck)
```

Output:
```
False
True
```

<!-- grade_id: cell-e963245866171c87, points: 0, grade: false, solution: true -->
```python
# some goes for flush

def extract_suit(card):
  return card[-1] # this is easy enough to udnerstand


def has_flush(hand):
  # Some code jsut changing the indice
  suits = [extract_suit(card) for card in hand];
  value,count = np.unique(suits,return_counts=True)
  return np.any(count >= 5)
```

<!-- grade_id: hw10_1b_2, points: 5, grade: true, solution: false -->
```python
print(has_flush(test_hand_1))
print(has_flush(test_hand_4))

assert not has_flush(test_hand_1)
assert has_flush(test_hand_4)

#
```

Output:
```
False
True
```

### Part C (7 points)

Using the functions you wrote above, use Monte Carlo simulation to calculate the probability of each of the following outcomes:

- A hand of 12 cards has 4 cards with the same rank  ("four-of-a-kind".)  __Give your answer in the variable `p_4kind`.__
- A hand of 12 cards has 5 cards with the same suit ("flush".)  __Give your answer in the variable `p_flush`.__

Your answers should be __accurate to within 0.01__ (i.e. 1% probability) - use enough hands to ensure this is true.  (Run a few trials to get a sense for the size of fluctuations, or estimate the standard error.)

<!-- grade_id: cell-c10760120558e2d1, points: 0, grade: false, solution: true -->
```python
#
N = 1000 # This is the number to simulate I will try to change it to test

four_kind_count = 0
flush_count = 0

for _ in range(N):
  hand = np.random.choice(deck, 12, replace=False) #
  #print(hand)

  if has_4kind(hand):
    four_kind_count += 1
  if has_flush(hand):
    flush_count += 1

p_4kind = four_kind_count / N
p_flush = flush_count / N

print(p_4kind)
print(p_flush)

```

Output:
```
0.031
0.495
```

**Design** and implement some tests of your function in the cell below. Explain the tests in comments, and verify that your code passes your tests.

<!-- grade_id: hw9_t1, points: 2, grade: true, solution: true -->
```python
# first thing I want to make my life easier, extract rnaks and suit and build it
# from that
ranks = [extract_rank(card) for card in deck];
suits = [extract_suit(card) for card in deck];

# make the suits unique
suits = np.unique(suits)
# print(suits)
#
# The below is only flush test
# Test 0:  hand with a flush
# expected treturn true
hand_flush = [r + suits[0] for r in ranks[:6]] + [r + suits[2] for r in ranks[6:12]]
# print (hand_flush)
assert has_flush(hand_flush) == True

# Test 1: Hand with 4 of a kind
hand_4kind = ["4" + suits[0] for i in range(4)] + [r + suits[0] for r in ranks[4:8]] + [r + suits[0] for r in ranks[8:12]]
assert has_4kind(hand_4kind) == True

# Test 2: This is a test I like four of all suits to see if the flush happen or not
hand_no_flush = [r + suits[0] for r in ranks[:3]] + [r + suits[1] for r in ranks[3:6]] + [r + suits[2] for r in ranks[6:9]] + [r + suits[3] for r in ranks[9:12]]
print(hand_no_flush)
assert has_flush(hand_no_flush) == False
#
#

```

Output:
```
['2♠', '3♠', '4♠', '5♣', '6♣', '7♣', '8♥', '9♥', '10♥', 'J♦', 'Q♦', 'K♦']
```

<!-- grade_id: hw10_1c, points: 5, grade: true, solution: false -->
```python
# testing cell
assert abs(p_4kind - 0.024) <= 0.02
assert abs(p_flush - 0.49) <= 0.02
```

## 10.2 - Monte Carlo and moments of inertia (22 points)

A sphere of radius 6 cm is made of solid resin, but with a cylindrical gap of radius $r=$ 2 cm and height $h=4$ cm centered at $\mathbf{R} = (2,2,2)$ cm from the origin of the larger sphere:

<img src="https://raw.githubusercontent.com/wlough/CU-Phys2600-Fall2025/main/lectures/img/sphere-cylinder2.png" style="max-width:100%" />

A 3D printer is probably the only way to fabricate something like this, so we'll assume the solid part of the sphere is made from plastic resin (density $\rho \sim 1.2\ \rm{g}/\rm{cm}^3$).  We would like to know the moment of inertia $I_z$ about the $z$-axis, given by the integral

$$
I_z = \int dV\ \rho(\mathbf{r}) r^2
$$
where $r$ is the cylindrical radius, i.e. the distance to the $z$-axis.  (Although you may not have seen this integral before, you should at least remember from intro physics that the moment of inertia acts like the mass for rotational motion: if we calculate $I_z$, then for spinning about the $z$-axis the _angular momentum_ $L$ will be related to the _angular velocity_ $\omega_z$ by $L = I_z \omega_z$.)



This will be a really frustrating integral to set up in any coordinate system, because of the off-center missing spherical region.  Fortunately, it's easy to do with Monte Carlo!  We'll do hit-or-miss, defining the __indicator function__

$$
f(x,y,z) = \begin{cases}
1, & \textrm{(x,y,z) inside object;} \\
0, & \textrm{(x,y,z) outside object.}
\end{cases}
$$

Then $\rho(\mathbf{r}) = \rho \cdot f(x,y,z)$, where $\rho$ is now just a constant.  Then the integral becomes

$$
I_z = \rho \int dV\ (x^2 + y^2)\ f(x,y,z) \\
= \rho V \langle (x^2 + y^2)\ f(x,y,z) \rangle_V \\
= \frac{\rho V}{N} \sum_{i=1}^N (x_i^2 + y_i^2)\ f(x_i,y_i,z_i)
$$

drawing $N$ points randomly in the three-volume $V$ containing the object.  To sum up, we can deal with this by doing hit-or-miss with a slight modification:

1. Draw points randomly in a three-dimensional cube surrounding the object.
2. Keep a running total over all points, adding $(x^2 + y^2)$ if it is _inside_ the object, and $0$ if it is _outside_.
3. Plug in to the formula above to get $I_z$, dividing by the __total__ number of points $N$ drawn.

Let's implement this in pieces.  We'll use NumPy arrays for everything, since we'd like to be able to efficiently calculate with lots of random points at once.

### Part A (5 points)

To make sure we set things up properly, let's begin with a case where we know the answer: a completely solid sphere.  Let's start by implementing the indicator function $f(x,y,z)$ separately.  In the cell below, __implement the function `f_sphere`__.  This function takes an $N \times 3$ NumPy array `r` as input, and produces a length-N mask which is True for any points in `r` which are inside the sphere.



<!-- grade_id: cell-c614d3edf25defee, points: 0, grade: false, solution: true -->
```python
def f_sphere(r, R=6):
    """
    Indicator function for the solid sphere of radius R.

    Args:
      - r: (Nx3) NumPy array, encoding N three-dimensional points in
      Cartesian coordinates.
      - R: optional, radius of the sphere to be tested against.
      Assumed to be centered at the origin.  Default is 6.

    Returns:
      - mask: a length-N NumPy mask, which is True for each point in r
      that is inside the radius-R sphere, and False for each point
      outside the sphere.

    """

    #
    return np.linalg.norm(r, axis=1) <= R
```

Design and implement some tests of your function in the cell below. Explain the tests in comments, and verify that your code passes your tests.

<!-- grade_id: hw9_t1, points: 2, grade: true, solution: true -->
```python
# I will write using the np assert method since I want the practice

test_pt = np.array([
    [1,1,1], # this is inside and should return true since sqrt(3) is less than 6 #True
    [2,2,2],# False since this is sqrt(12) < 6 it's not less that                        #False
    [-2,2,0], # This is inside since this is just 8                               #True
    [10, 0, 0 ] # outside False                                                   #Fasle
])

import numpy.testing as npt

npt.assert_allclose(f_sphere(test_pt), [True, True, True, False])
```

<!-- grade_id: hw10_2a, points: 3, grade: true, solution: false -->
```python
import numpy.testing as npt

test_pts = np.array(
    [
        [0, 0, 0],  # Inside
        [-1, -1, -1],  # Inside
        [-3, -3, -4],  # Inside
        [6, 6, 6],  # Outside
        [2, 2, 2],  # Inside
    ]
)

print(f_sphere(test_pts))

import numpy.testing as npt

npt.assert_allclose(f_sphere(test_pts), [True, True, True, False, True])

assert not f_sphere(np.array([[-3, -3, -4]]), R=1)
assert not f_sphere(np.array([[-1, -1, -1]]), R=1)

```

Output:
```
[ True  True  True False  True]
```

### Part B (6 points)

Now we'll implement the rest of the code to calculate $I_z$ using Monte Carlo integration.  When we move on to do the more complicated shape drawn above, the only difference will be changing out the indicator function - everything else will be the same for the solid sphere.  So, let's write a single function `find_Iz` that takes the indicator function as an argument.

From here forward, we'll hard-code to our specific problem of a sphere of radius $R = 6$ cm.  Thus, when doing our hit-or-miss algorithm we should draw random points from the cube $-6 \leq x,y,z \leq 6$.




<!-- grade_id: cell-5397c64c7abb3601, points: 0, grade: false, solution: true -->
```python
def find_Iz(N, f):
    """
    Compute the moment of inertia about the z-axis for a spherical object
    of radius 6 cm.

    Args:
      - N: number of random points to use.
      - f: indicator function (see part A.)

    Returns:
      - I_z: float, the moment of inertia of the object.

    """

    #
    R = 6
    rho = 0.0012
    r = (np.random.rand(N, 3) - 0.5) * (2 * R)

    x = r[:, 0]
    y = r[:,1]

    inside = f(r)

    intergrand  = (x**2 + y**2) * inside
    V_cube = (2*R)**3

    I_z = rho * V_cube * np.mean(intergrand)

    return I_z
```

Now let's test - __use `find_Iz` and `f_sphere`__ to calculate $I_z$ for the solid sphere in the cell below.  This will be a good test, since we know the answers analytically.  The expected result for a solid sphere is:

$$
I_z = \frac{2}{5} m r^2 \\
m = \frac{4}{3} \pi r^3 \rho
$$

which for $r = 6$ cm and $\rho = 1.2\ \rm{g}/\rm{cm}^3$ gives $I_z = 15.63\ \rm{kg} \cdot \rm{cm}^2$.  __Save your answer to the variable `Iz_solid`.__

(You should try to match your estimate for $I_z$ to __within 5%__ of the exact answer - adjust the number of random points you're using as needed.)

<!-- grade_id: cell-d74180196317ff31, points: 0, grade: false, solution: true -->
```python
#
N = 3000000
Iz_solid = find_Iz(N, f_sphere)
print(Iz_solid)
```

Output:
```
15.6138333465499
```

<!-- grade_id: hw10_2b, points: 6, grade: true, solution: false -->
```python
print(f"Iz (solid sphere, exact) = {15.63}")
print(f"Iz (solid sphere, Monte Carlo) = {Iz_solid}")

assert np.abs(Iz_solid - 15.63) / 15.63 <= 0.05
```

Output:
```
Iz (solid sphere, exact) = 15.63
Iz (solid sphere, Monte Carlo) = 15.6138333465499
```

### Part C (8 points)

Now on to the more complicated object.  __Implement the indicator function `f_object`__ in the cell below, then run the test cell to make sure it's working for a few example points.  This should be implemented to work on _arrays_, as the docstring indicates (and as you did in part A.)

<!-- grade_id: cell-431568016bcaae4a, points: 0, grade: false, solution: true -->
```python
def f_object(r):
    """
    Indicator function for the sphere with a cylindrical hole.

    Input: r, an (Nx3) NumPy array, where each row contains three numbers (x,y,z),
    corresponding to a point in space.  All lengths are taken to be in cm.

    Output: m, a length-N NumPy mask, which is True if (x,y,z) is inside the solid
    part of the compound object described above, and False if (x,y,z) is outside it.
    """

    x = r[:, 0]
    y = r[:, 1]
    z = r[:, 2]


    # now we need to find the inside of sphere we know the radius 6
    inside_sphere = (x*x+y*y+z*z) <= 6*6

    c_center = np.array([2,2,2])
    c_radius = 2
    c_height = 4

    in_c_r = (x - c_center[0])**2 + (y - c_center[1])**2 <= c_radius**2

    in_c_v = (z>= (c_center[-1] - c_height/2)) & (z <= (c_center[-1] + c_height/2))

    # Finally
    in_cy = in_c_r & in_c_v

    return inside_sphere & (~in_cy)
#



```

Design and implement some tests of your function in the cell below. Explain the tests in comments, and verify that your code passes your tests.

<!-- grade_id: hw9_t1, points: 2, grade: true, solution: true -->
```python
# Here I want to change I test; I want to test 100 random points and see
# helper function to produce

def generate_inside(N=100):
  pts_test = []

  for _ in range(N):
    x = np.random.uniform(-6,6)
    y = np.random.uniform(-6,6)
    z = np.random.uniform(-6,6)

    inside_sphere = (x*x+y*y+z*z) <= 6*6
    in_rad = (x-2)**2 + (y-2)*2 <=4
    in_vert = (z>= 0) and (z <= 4)
    inside_cyl = in_rad & in_vert

    if (inside_sphere and not inside_cyl):
      pts_test.append([x,y,z])

  return np.array(pts_test)

# Hole
def generate_inside_hole(N=100):
  pts_test = []

  for _ in range(N):
    x = np.random.uniform(0,4)
    y = np.random.uniform(0,4)
    z = np.random.uniform(0,4)

    in_rad = (x-2)**2 + (y-2)*2 <=4
    in_vert = (z>= 0) and (z <= 4)

    if (in_rad and in_vert):
      pts_test.append([x,y,z])

  return np.array(pts_test)


def generate_outside(N=100):
  pts_test = []

  for _ in range(N):
    x = np.random.uniform(-6,6)
    y = np.random.uniform(-6,6)
    z = np.random.uniform(-6,6)

    if (x*x+y*y+z*z > 6*6):
      pts_test.append([x,y,z])

  return np.array(pts_test)

# Now I can test

inside_obj = generate_inside(50)
outside_obj = generate_outside(50)
inside_hole = generate_inside_hole(50)

#
assert np.all((f_object(inside_obj)))  # all of this should return true
assert not np.all((f_object(inside_hole)))
assert not np.all((f_object(outside_obj)))

print("Tests Done All passed, Coding is workign and some how fun!!")
#


```

Output:
```
Tests Done All passed, Coding is workign and some how fun!!
```

<!-- grade_id: hw10_2c, points: 6, grade: true, solution: false -->
```python

test_pts = np.array(
    [
        [0, 0, 0],  # Inside
        [-1, -1, -1],  # Inside
        [-3, -3, -4],  # Inside
        [5, 0, -2],  # Inside
        [6, 6, 6],  # Outside
        [-2, 2, 2],  # Inside
        [2, 2, 3],  # Outside (in the gap)
        [3, 4, 3],  # Inside
        [0.1, 0.1, 0.5],  # Inside
    ]
)

print(f_object(test_pts))

import numpy.testing as npt

npt.assert_allclose( # 7 True # 2 False
    f_object(test_pts), [True, True, True, True, False, True, False, True, True]
)
```

Output:
```
[ True  True  True  True False  True False  True  True]
```

### Part D (3 points)

Now put it all together: use `find_Iz` and `f_object` to __find the moment of inertia__ for the sphere-missing-a-sphere, setting `N` large enough that you can trust the result to within 0.1.  __About how much smaller is the full $I_z$ than it was for the solid sphere, in percent?__  (Answer using a comment.)

<!-- grade_id: hw10_2d, points: 3, grade: true, solution: true -->
```python
# Now is the end of this, we just need to combine the function togather and simulate
N = 800_000 # First I will test with 200_000

Iz_solid = find_Iz(N, f_sphere)
Iz_object = find_Iz(N, f_object)

# printing to see
print(f"Iz (solid sphere, Monte Carlo) = {Iz_solid}")
print(f"Iz (object, Monte Carlo) = {Iz_object}")

# fidning the percent
PS = (Iz_solid-Iz_object) / Iz_solid*100
print(f"Percent smaller = {PS}%")

#it's gonname 3% to 4%  The max I got is like 4.5%
# if we increase N it will be between 3% and max as 5%
```

Output:
```
Iz (solid sphere, Monte Carlo) = 15.630062912467745
Iz (object, Monte Carlo) = 15.017600803704099
Percent smaller = 3.918487802599303%
```